(assume a0 (! (forall ((?v0 S_v_tm$) (?v1 S_v_tm$)) (= (fun_app$ (fun_app$a uu$ ?v0) ?v1) (and (= (vars$ ?v1) bot$) (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))))))))) :named @p_0))
(assume a1 (! (forall ((?v0 S_v_tm$) (?v1 S_v_tm$)) (= (fun_app$ (fun_app$a uua$ ?v0) ?v1) (and (= (vars$ ?v1) bot$) (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) (! (gt$ ground_heads_var$ gt_sym$ delta$ extf$ arity_sym$ wt_sym$ coef_sym$) :named @p_2)) (args$ ?v2)) (args$ ?v3))))))))))))) :named @p_1))
(assume a2 (! (forall ((?v0 V_pvar_hmultiset_tpoly$) (?v1 V_pvar_hmultiset_tpoly$) (?v2 V_pvar_hmultiset_tpoly$)) (=> (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v0 ?v1) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v1 ?v2)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v0 ?v2))) :named @p_3))
(assume a3 (! (forall ((?v0 V_pvar_hmultiset_tpoly$) (?v1 V_pvar_hmultiset_tpoly$) (?v2 V_pvar_hmultiset_tpoly$)) (=> (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v0 ?v1) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v1 ?v2)) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v0 ?v2))) :named @p_4))
(assume a4 (! (forall ((?v0 S$) (?v1 S_bool_fun$)) (= (member$ ?v0 (collect$ ?v1)) (fun_app$f ?v1 ?v0))) :named @p_5))
(assume a5 (! (forall ((?v0 S_v_tm_s_v_tm_prod$) (?v1 S_v_tm_s_v_tm_prod_bool_fun$)) (= (member$a ?v0 (collect$a ?v1)) (fun_app$g ?v1 ?v0))) :named @p_6))
(assume a6 (! (forall ((?v0 S_v_tm_s_v_tm_prod$) (?v1 S_v_tm_s_v_tm_prod_set$) (?v2 S_v_tm_s_v_tm_prod_set$)) (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) (! false :named @p_8)))) @p_8)) :named @p_7))
(assume a7 (! (forall ((?v0 S_v_tm_s_v_tm_bool_fun_fun$) (?v1 S_v_tm$) (?v2 S_v_tm$)) (= (fun_app$g (case_prod$ ?v0) (pair$ ?v1 ?v2)) (fun_app$ (fun_app$a ?v0 ?v1) ?v2))) :named @p_9))
(assume a8 (! (forall ((?v0 S_set$) (?v1 S_set$)) (=> (forall ((?v2 S$)) (=> (member$ ?v2 ?v0) (member$ ?v2 ?v1))) (less_eq$ ?v0 ?v1))) :named @p_10))
(assume a9 (! (forall ((?v0 S_v_tm_s_v_tm_prod_set$) (?v1 S_v_tm_s_v_tm_prod_set$)) (=> (forall ((?v2 S_v_tm_s_v_tm_prod$)) (=> (member$a ?v2 ?v0) (member$a ?v2 ?v1))) (less_eq$a ?v0 ?v1))) :named @p_11))
(assume a10 (! (not (! (less_eq$a (! (relcomp$ (! (collect$a (! (case_prod$ uu$) :named @p_16)) :named @p_15) (! (collect$a (! (case_prod$ uua$) :named @p_18)) :named @p_17)) :named @p_14) @p_15) :named @p_13)) :named @p_12))
(step t0 (cl (! (=> (! (forall ((?v0 V_pvar_hmultiset_tpoly$) (?v1 V_pvar_hmultiset_tpoly$) (?v2 V_pvar_hmultiset_tpoly$)) (or (not (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v0 ?v1)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v1 ?v2)) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v0 ?v2))) :named @p_20) (! (or (! (not (! (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (! (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ (! (choice ((?v5 S_v_tm$)) (not (or (not (= (pair$ (! (choice ((?v3 S_v_tm$)) (not (forall ((?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= (pair$ ?v3 ?v5) (! (choice ((?v2 S_v_tm_s_v_tm_prod$)) (not (or (not (member$a ?v2 @p_14)) (member$a ?v2 @p_15)))) :named @p_27))) (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17)))))) :named @p_26) ?v5) @p_27)) (! (not (! (member$a (! (pair$ @p_26 (! (choice ((?v4 S_v_tm$)) (not (forall ((?v5 S_v_tm$)) (or (not (= (pair$ @p_26 ?v5) @p_27)) (not (member$a (pair$ @p_26 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17)))))) :named @p_31)) :named @p_30) @p_15) :named @p_29)) :named @p_28) (not (member$a (pair$ @p_31 ?v5) @p_17))))) :named @p_25)) :named @p_24) (! (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ @p_31) :named @p_32)) :named @p_23)) :named @p_22) (! (not (! (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ @p_32 (! (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ @p_26) :named @p_35)) :named @p_34)) :named @p_33) (! (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ @p_24 @p_35) :named @p_36)) :named @p_21)) :named @p_19) @p_20) :rule implies_neg1)
(anchor :step t1)
(assume t1.a0 @p_20)
(step t1.t0 (cl (or (! (not @p_20) :named @p_37) @p_21)) :rule forall_inst :args (@p_24 @p_32 @p_35))
(step t1.t1 (cl @p_37 @p_21) :rule or :premises (t1.t0))
(step t1.t2 (cl (! (not (! (= @p_4 @p_20) :named @p_39)) :named @p_38) (! (not @p_4) :named @p_40) @p_20) :rule equiv_pos2)
(step t1.t3 (cl @p_39) :rule hole)
(step t1.t4 (cl @p_20) :rule resolution :premises (t1.t2 t1.t3 a3) :args (@p_39 @p_8 @p_4 @p_8))
(step t1.t5 (cl @p_21) :rule resolution :premises (t1.t1 t1.t4) :args (@p_20 @p_8))
(step t1 (cl @p_37 @p_21) :rule subproof :discharge (t1.a0))
(step t2 (cl @p_19 @p_21) :rule resolution :premises (t0 t1) :args (@p_20 (! true :named @p_41)))
(step t3 (cl @p_19 (! (not @p_21) :named @p_42)) :rule implies_neg2)
(step t4 (cl @p_19 @p_19) :rule resolution :premises (t2 t3) :args (@p_21 @p_41))
(step t5 (cl @p_19) :rule contraction :premises (t4))
(step t6 (cl @p_42 @p_22 @p_33 @p_36) :rule or_pos)
(step t7 (cl (! (not (! (= (! (fun_app$ (! (fun_app$a uu$ @p_26) :named @p_46) @p_25) :named @p_45) (! (and (! (= bot$ (! (vars$ @p_25) :named @p_49)) :named @p_48) (! (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ @p_24 @p_35) :named @p_50) @p_36 (! (gt_hd$ gt_sym$ ground_heads_var$ (! (head$ @p_25) :named @p_52) (! (head$ @p_26) :named @p_53)) :named @p_51)) :named @p_47)) :named @p_44)) :named @p_43) @p_45 (! (not @p_47) :named @p_54)) :rule equiv_pos1)
(step t8 (cl (! (not (! (= (! (fun_app$g @p_16 (! (pair$ @p_26 @p_25) :named @p_58)) :named @p_57) @p_45) :named @p_56)) :named @p_55) @p_57 (! (not @p_45) :named @p_59)) :rule equiv_pos1)
(step t9 (cl (= (! (= @p_57 @p_8) :named @p_60) (! (not @p_57) :named @p_61))) :rule equiv_simplify)
(step t10 (cl (= (! (or (! (not (! (= @p_27 @p_58) :named @p_64)) :named @p_63) @p_28 (! (not (! (member$a (! (pair$ @p_31 @p_25) :named @p_67) @p_17) :named @p_66)) :named @p_65)) :named @p_62) @p_62)) :rule refl)
(step t11 (cl (= (! (not @p_63) :named @p_68) @p_64)) :rule hole)
(step t12 (cl (not (! (= (! (or @p_62 @p_68) :named @p_70) (! (or @p_62 @p_64) :named @p_71)) :named @p_69)) (not @p_70) @p_71) :rule equiv_pos2)
(step t13 (cl @p_69) :rule cong :premises (t10 t11))
(step t14 (cl @p_62 @p_68) :rule or_neg :args (0))
(step t15 (cl @p_70 (! (not @p_62) :named @p_72)) :rule or_neg :args (0))
(step t16 (cl @p_70 (not @p_68)) :rule or_neg :args (1))
(step t17 (cl @p_70 @p_70) :rule resolution :premises (t14 t15 t16) :args (@p_62 @p_41 @p_68 @p_41))
(step t18 (cl @p_70) :rule contraction :premises (t17))
(step t19 (cl @p_71) :rule resolution :premises (t12 t13 t18) :args (@p_69 @p_8 @p_70 @p_8))
(step t20 (cl @p_62 @p_64) :rule or :premises (t19))
(step t21 (cl (= (! (not (! (not (! (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= (pair$ ?v3 ?v5) @p_27)) (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17)))) :named @p_75)) :named @p_74)) :named @p_73) @p_75)) :rule hole)
(step t22 (cl (= @p_72 @p_72)) :rule refl)
(step t23 (cl (= @p_74 @p_74)) :rule refl)
(step t24 (cl (= (! (= @p_58 @p_27) :named @p_76) @p_64)) :rule hole)
(step t25 (cl (= (! (not @p_76) :named @p_77) @p_63)) :rule cong :premises (t24))
(step t26 (cl (= @p_28 @p_28)) :rule refl)
(step t27 (cl (= @p_65 @p_65)) :rule refl)
(step t28 (cl (= (! (or @p_77 @p_28 @p_65) :named @p_78) @p_62)) :rule cong :premises (t25 t26 t27))
(step t29 (cl (= (! (not @p_78) :named @p_79) @p_72)) :rule cong :premises (t28))
(step t30 (cl (not (! (= (! (=> @p_74 @p_79) :named @p_81) (! (=> @p_74 @p_72) :named @p_82)) :named @p_80)) (not @p_81) @p_82) :rule equiv_pos2)
(step t31 (cl @p_80) :rule cong :premises (t23 t29))
(step t32 (cl @p_81 @p_74) :rule implies_neg1)
(anchor :step t33)
(assume t33.a0 @p_74)
(anchor :step t33.t0 :args ((:= (?v3 S_v_tm$) @p_26) (:= (?v4 S_v_tm$) @p_31) (:= (?v5 S_v_tm$) @p_25)))
(step t33.t0.t0 (cl (= (or (not (= (pair$ ?v3 ?v5) @p_27)) (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17))) @p_78)) :rule refl)
(step t33.t0 (cl (= @p_75 @p_78)) :rule sko_forall)
(step t33.t1 (cl (! (not (! (or (! (not (! (member$a @p_27 @p_14) :named @p_86)) :named @p_85) @p_74) :named @p_84)) :named @p_83) @p_85 @p_74) :rule or_pos)
(step t33.t2 (cl (! (= (! (or @p_85 (! (member$a @p_27 @p_15) :named @p_89)) :named @p_88) @p_88) :named @p_87)) :rule refl)
(step t33.t3 (cl (! (= (! (not @p_85) :named @p_91) @p_86) :named @p_90)) :rule hole)
(step t33.t4 (cl (! (not (! (= (! (or @p_88 @p_91) :named @p_94) (! (or @p_88 @p_86) :named @p_95)) :named @p_93)) :named @p_92) (! (not @p_94) :named @p_96) @p_95) :rule equiv_pos2)
(step t33.t5 (cl @p_93) :rule cong :premises (t33.t2 t33.t3))
(step t33.t6 (cl @p_88 @p_91) :rule or_neg :args (0))
(step t33.t7 (cl @p_94 (! (not @p_88) :named @p_97)) :rule or_neg :args (0))
(step t33.t8 (cl @p_94 (! (not @p_91) :named @p_98)) :rule or_neg :args (1))
(step t33.t9 (cl @p_94 @p_94) :rule resolution :premises (t33.t6 t33.t7 t33.t8) :args (@p_88 @p_41 @p_91 @p_41))
(step t33.t10 (cl @p_94) :rule contraction :premises (t33.t9))
(step t33.t11 (cl @p_95) :rule resolution :premises (t33.t4 t33.t5 t33.t10) :args (@p_93 @p_8 @p_94 @p_8))
(step t33.t12 (cl @p_88 @p_86) :rule or :premises (t33.t11))
(anchor :step t33.t13 :args ((:= (?v2 S_v_tm_s_v_tm_prod$) @p_27)))
(step t33.t13.t0 (cl (= (or (not (member$a ?v2 @p_14)) (member$a ?v2 @p_15)) @p_88)) :rule refl)
(step t33.t13 (cl (! (= (! (forall ((?v2 S_v_tm_s_v_tm_prod$)) (or (not (member$a ?v2 @p_14)) (member$a ?v2 @p_15))) :named @p_100) @p_88) :named @p_99)) :rule sko_forall)
(step t33.t14 (cl (! (not (! (or (! (not @p_100) :named @p_103) @p_13) :named @p_102)) :named @p_101) @p_103 @p_13) :rule or_pos)
(step t33.t15 (cl (! (or (! (not (! (forall ((?v0 S_v_tm_s_v_tm_prod_set$) (?v1 S_v_tm_s_v_tm_prod_set$)) (or (not (forall ((?v2 S_v_tm_s_v_tm_prod$)) (or (not (member$a ?v2 ?v0)) (member$a ?v2 ?v1)))) (less_eq$a ?v0 ?v1))) :named @p_106)) :named @p_105) @p_102) :named @p_104)) :rule forall_inst :args (@p_14 @p_15))
(anchor :step t33.t16 :args ((?v0 S_v_tm_s_v_tm_prod_set$) (?v1 S_v_tm_s_v_tm_prod_set$) (:= (?v0 S_v_tm_s_v_tm_prod_set$) ?v0) (:= (?v1 S_v_tm_s_v_tm_prod_set$) ?v1)))
(step t33.t16.t0 (cl (= (forall ((?v2 S_v_tm_s_v_tm_prod$)) (=> (member$a ?v2 ?v0) (member$a ?v2 ?v1))) (forall ((?v2 S_v_tm_s_v_tm_prod$)) (or (not (member$a ?v2 ?v0)) (member$a ?v2 ?v1))))) :rule hole)
(step t33.t16.t1 (cl (= (less_eq$a ?v0 ?v1) (less_eq$a ?v0 ?v1))) :rule refl)
(step t33.t16.t2 (cl (= (=> (forall ((?v2 S_v_tm_s_v_tm_prod$)) (=> (member$a ?v2 ?v0) (member$a ?v2 ?v1))) (less_eq$a ?v0 ?v1)) (=> (forall ((?v2 S_v_tm_s_v_tm_prod$)) (or (not (member$a ?v2 ?v0)) (member$a ?v2 ?v1))) (less_eq$a ?v0 ?v1)))) :rule cong :premises (t33.t16.t0 t33.t16.t1))
(step t33.t16 (cl (! (= @p_11 (! (forall ((?v0 S_v_tm_s_v_tm_prod_set$) (?v1 S_v_tm_s_v_tm_prod_set$)) (=> (forall ((?v2 S_v_tm_s_v_tm_prod$)) (or (not (member$a ?v2 ?v0)) (member$a ?v2 ?v1))) (less_eq$a ?v0 ?v1))) :named @p_108)) :named @p_107)) :rule bind)
(step t33.t17 (cl (! (= @p_108 @p_106) :named @p_109)) :rule hole)
(step t33.t18 (cl (! (= (! (forall ((?v0 S_v_tm_s_v_tm_prod$) (?v1 S_v_tm_s_v_tm_prod_set$) (?v2 S_v_tm_s_v_tm_prod_set$)) (or (not (member$a ?v0 (relcomp$ ?v1 ?v2))) (not (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2))))))) :named @p_111) @p_111) :named @p_110)) :rule refl)
(step t33.t19 (cl (! (= @p_85 @p_85) :named @p_112)) :rule refl)
(anchor :step t33.t20 :args ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$) (:= (?v3 S_v_tm$) ?v3) (:= (?v4 S_v_tm$) ?v4) (:= (?v5 S_v_tm$) ?v5)))
(step t33.t20.t0 (cl (= (= @p_27 (pair$ ?v3 ?v5)) (= (pair$ ?v3 ?v5) @p_27))) :rule hole)
(step t33.t20.t1 (cl (= (not (= @p_27 (pair$ ?v3 ?v5))) (not (= (pair$ ?v3 ?v5) @p_27)))) :rule cong :premises (t33.t20.t0))
(step t33.t20.t2 (cl (= (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v3 ?v4) @p_15)))) :rule refl)
(step t33.t20.t3 (cl (= (not (member$a (pair$ ?v4 ?v5) @p_17)) (not (member$a (pair$ ?v4 ?v5) @p_17)))) :rule refl)
(step t33.t20.t4 (cl (= (or (not (= @p_27 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17))) (or (not (= (pair$ ?v3 ?v5) @p_27)) (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17))))) :rule cong :premises (t33.t20.t1 t33.t20.t2 t33.t20.t3))
(step t33.t20 (cl (! (= (! (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= @p_27 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17)))) :named @p_114) @p_75) :named @p_113)) :rule bind)
(step t33.t21 (cl (! (= (! (not @p_114) :named @p_116) @p_74) :named @p_115)) :rule cong :premises (t33.t20))
(step t33.t22 (cl (! (= (! (or @p_85 @p_116) :named @p_118) @p_84) :named @p_117)) :rule cong :premises (t33.t19 t33.t21))
(step t33.t23 (cl (! (not (! (= (! (=> @p_111 @p_118) :named @p_121) (! (=> @p_111 @p_84) :named @p_122)) :named @p_120)) :named @p_119) (! (not @p_121) :named @p_123) @p_122) :rule equiv_pos2)
(step t33.t24 (cl @p_120) :rule cong :premises (t33.t18 t33.t22))
(step t33.t25 (cl @p_121 @p_111) :rule implies_neg1)
(anchor :step t33.t26)
(assume t33.t26.a0 @p_111)
(step t33.t26.t0 (cl (! (or (! (not @p_111) :named @p_125) @p_118) :named @p_124)) :rule forall_inst :args (@p_27 @p_15 @p_17))
(anchor :step t33.t26.t1 :args ((?v0 S_v_tm_s_v_tm_prod$) (?v1 S_v_tm_s_v_tm_prod_set$) (?v2 S_v_tm_s_v_tm_prod_set$) (:= (?v0 S_v_tm_s_v_tm_prod$) ?v0) (:= (?v1 S_v_tm_s_v_tm_prod_set$) ?v1) (:= (?v2 S_v_tm_s_v_tm_prod_set$) ?v2)))
(step t33.t26.t1.t0 (cl (= (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) @p_8) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)))))) :rule hole)
(step t33.t26.t1.t1 (cl (= (member$a ?v0 (relcomp$ ?v1 ?v2)) (member$a ?v0 (relcomp$ ?v1 ?v2)))) :rule refl)
(anchor :step t33.t26.t1.t2 :args ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$) (:= (?v3 S_v_tm$) ?v3) (:= (?v4 S_v_tm$) ?v4) (:= (?v5 S_v_tm$) ?v5)))
(step t33.t26.t1.t2.t0 (cl (= (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8) (not (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule hole)
(step t33.t26.t1.t2.t1 (cl (= (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) :rule hole)
(step t33.t26.t1.t2.t2 (cl (= (not (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))))) :rule cong :premises (t33.t26.t1.t2.t1))
(step t33.t26.t1.t2.t3 (cl (= (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))))) :rule trans :premises (t33.t26.t1.t2.t0 t33.t26.t1.t2.t2))
(step t33.t26.t1.t2 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule bind)
(step t33.t26.t1.t3 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule hole)
(step t33.t26.t1.t4 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule trans :premises (t33.t26.t1.t2 t33.t26.t1.t3))
(step t33.t26.t1.t5 (cl (= (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2))))))) :rule cong :premises (t33.t26.t1.t1 t33.t26.t1.t4))
(step t33.t26.t1.t6 (cl (= (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)))) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))))) :rule cong :premises (t33.t26.t1.t5))
(step t33.t26.t1.t7 (cl (= (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) @p_8) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))))) :rule trans :premises (t33.t26.t1.t0 t33.t26.t1.t6))
(step t33.t26.t1 (cl (! (= @p_7 (! (forall ((?v0 S_v_tm_s_v_tm_prod$) (?v1 S_v_tm_s_v_tm_prod_set$) (?v2 S_v_tm_s_v_tm_prod_set$)) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2))))))) :named @p_127)) :named @p_126)) :rule bind)
(step t33.t26.t2 (cl (! (= @p_127 @p_111) :named @p_128)) :rule hole)
(step t33.t26.t3 (cl @p_125 @p_118) :rule or :premises (t33.t26.t0))
(step t33.t26.t4 (cl (! (not (! (= @p_7 @p_111) :named @p_130)) :named @p_129) (! (not @p_7) :named @p_131) @p_111) :rule equiv_pos2)
(step t33.t26.t5 (cl @p_130) :rule trans :premises (t33.t26.t1 t33.t26.t2))
(step t33.t26.t6 (cl @p_111) :rule resolution :premises (t33.t26.t4 t33.t26.t5 a6) :args (@p_130 @p_8 @p_7 @p_8))
(step t33.t26.t7 (cl @p_118) :rule resolution :premises (t33.t26.t3 t33.t26.t6) :args (@p_111 @p_8))
(step t33.t26 (cl @p_125 @p_118) :rule subproof :discharge (t33.t26.a0))
(step t33.t27 (cl @p_121 @p_118) :rule resolution :premises (t33.t25 t33.t26) :args (@p_111 @p_41))
(step t33.t28 (cl @p_121 (! (not @p_118) :named @p_132)) :rule implies_neg2)
(step t33.t29 (cl @p_121 @p_121) :rule resolution :premises (t33.t27 t33.t28) :args (@p_118 @p_41))
(step t33.t30 (cl @p_121) :rule contraction :premises (t33.t29))
(step t33.t31 (cl @p_122) :rule resolution :premises (t33.t23 t33.t24 t33.t30) :args (@p_120 @p_8 @p_121 @p_8))
(anchor :step t33.t32 :args ((?v0 S_v_tm_s_v_tm_prod$) (?v1 S_v_tm_s_v_tm_prod_set$) (?v2 S_v_tm_s_v_tm_prod_set$) (:= (?v0 S_v_tm_s_v_tm_prod$) ?v0) (:= (?v1 S_v_tm_s_v_tm_prod_set$) ?v1) (:= (?v2 S_v_tm_s_v_tm_prod_set$) ?v2)))
(step t33.t32.t0 (cl (= (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) @p_8) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)))))) :rule hole)
(step t33.t32.t1 (cl (= (member$a ?v0 (relcomp$ ?v1 ?v2)) (member$a ?v0 (relcomp$ ?v1 ?v2)))) :rule refl)
(anchor :step t33.t32.t2 :args ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$) (:= (?v3 S_v_tm$) ?v3) (:= (?v4 S_v_tm$) ?v4) (:= (?v5 S_v_tm$) ?v5)))
(step t33.t32.t2.t0 (cl (= (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8) (not (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule hole)
(step t33.t32.t2.t1 (cl (= (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) :rule hole)
(step t33.t32.t2.t2 (cl (= (not (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))))) :rule cong :premises (t33.t32.t2.t1))
(step t33.t32.t2.t3 (cl (= (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))))) :rule trans :premises (t33.t32.t2.t0 t33.t32.t2.t2))
(step t33.t32.t2 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule bind)
(step t33.t32.t3 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule hole)
(step t33.t32.t4 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule trans :premises (t33.t32.t2 t33.t32.t3))
(step t33.t32.t5 (cl (= (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2))))))) :rule cong :premises (t33.t32.t1 t33.t32.t4))
(step t33.t32.t6 (cl (= (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)))) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))))) :rule cong :premises (t33.t32.t5))
(step t33.t32.t7 (cl (= (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) @p_8) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))))) :rule trans :premises (t33.t32.t0 t33.t32.t6))
(step t33.t32 (cl @p_126) :rule bind)
(step t33.t33 (cl @p_128) :rule hole)
(step t33.t34 (cl (not (! (= @p_74 @p_79) :named @p_133)) @p_73 @p_79) :rule equiv_pos2)
(step t33.t35 (cl @p_133) :rule cong :premises (t33.t0))
(step t33.t36 (cl @p_85 @p_74 @p_83) :rule reordering :premises (t33.t1))
(step t33.t37 (cl @p_86 @p_88) :rule reordering :premises (t33.t12))
(step t33.t38 (cl (! (not (! (= @p_103 @p_97) :named @p_135)) :named @p_134) (! (not @p_103) :named @p_136) @p_97) :rule equiv_pos2)
(step t33.t39 (cl @p_135) :rule cong :premises (t33.t13))
(step t33.t40 (cl @p_13 @p_103 @p_101) :rule reordering :premises (t33.t14))
(step t33.t41 (cl @p_105 @p_102) :rule or :premises (t33.t15))
(step t33.t42 (cl (! (not (! (= @p_11 @p_106) :named @p_138)) :named @p_137) (! (not @p_11) :named @p_139) @p_106) :rule equiv_pos2)
(step t33.t43 (cl @p_138) :rule trans :premises (t33.t16 t33.t17))
(step t33.t44 (cl @p_106) :rule resolution :premises (t33.t42 t33.t43 a9) :args (@p_138 @p_8 @p_11 @p_8))
(step t33.t45 (cl @p_102) :rule resolution :premises (t33.t41 t33.t44) :args (@p_106 @p_8))
(step t33.t46 (cl @p_103) :rule resolution :premises (t33.t40 a10 t33.t45) :args (@p_13 @p_41 @p_102 @p_8))
(step t33.t47 (cl @p_97) :rule resolution :premises (t33.t38 t33.t39 t33.t46) :args (@p_135 @p_8 @p_103 @p_8))
(step t33.t48 (cl @p_86) :rule resolution :premises (t33.t37 t33.t47) :args (@p_88 @p_41))
(step t33.t49 (cl @p_125 @p_84) :rule implies :premises (t33.t31))
(step t33.t50 (cl @p_129 @p_131 @p_111) :rule equiv_pos2)
(step t33.t51 (cl @p_130) :rule trans :premises (t33.t32 t33.t33))
(step t33.t52 (cl @p_111) :rule resolution :premises (t33.t50 t33.t51 a6) :args (@p_130 @p_8 @p_7 @p_8))
(step t33.t53 (cl @p_84) :rule resolution :premises (t33.t49 t33.t52) :args (@p_111 @p_8))
(step t33.t54 (cl @p_74) :rule resolution :premises (t33.t36 t33.t48 t33.t53) :args (@p_86 @p_8 @p_84 @p_8))
(step t33.t55 (cl @p_79) :rule resolution :premises (t33.t34 t33.t35 t33.t54) :args (@p_133 @p_8 @p_74 @p_8))
(step t33 (cl @p_73 @p_79) :rule subproof :discharge (t33.a0))
(step t34 (cl @p_81 @p_79) :rule resolution :premises (t32 t33) :args (@p_74 @p_41))
(step t35 (cl @p_81 (not @p_79)) :rule implies_neg2)
(step t36 (cl @p_81 @p_81) :rule resolution :premises (t34 t35) :args (@p_79 @p_41))
(step t37 (cl @p_81) :rule contraction :premises (t36))
(step t38 (cl @p_82) :rule resolution :premises (t30 t31 t37) :args (@p_80 @p_8 @p_81 @p_8))
(step t39 (cl (not (! (= (! (or @p_73 @p_72) :named @p_141) (! (or @p_75 @p_72) :named @p_142)) :named @p_140)) (not @p_141) @p_142) :rule equiv_pos2)
(step t40 (cl @p_140) :rule cong :premises (t21 t22))
(step t41 (cl @p_73 @p_72) :rule implies :premises (t38))
(step t42 (cl @p_141 (not @p_73)) :rule or_neg :args (0))
(step t43 (cl @p_141 (not @p_72)) :rule or_neg :args (1))
(step t44 (cl @p_141 @p_141) :rule resolution :premises (t41 t42 t43) :args (@p_73 @p_41 @p_72 @p_41))
(step t45 (cl @p_141) :rule contraction :premises (t44))
(step t46 (cl @p_142) :rule resolution :premises (t39 t40 t45) :args (@p_140 @p_8 @p_141 @p_8))
(step t47 (cl @p_83 @p_85 @p_74) :rule or_pos)
(step t48 (cl @p_87) :rule refl)
(step t49 (cl @p_90) :rule hole)
(step t50 (cl @p_92 @p_96 @p_95) :rule equiv_pos2)
(step t51 (cl @p_93) :rule cong :premises (t48 t49))
(step t52 (cl @p_88 @p_91) :rule or_neg :args (0))
(step t53 (cl @p_94 @p_97) :rule or_neg :args (0))
(step t54 (cl @p_94 @p_98) :rule or_neg :args (1))
(step t55 (cl @p_94 @p_94) :rule resolution :premises (t52 t53 t54) :args (@p_88 @p_41 @p_91 @p_41))
(step t56 (cl @p_94) :rule contraction :premises (t55))
(step t57 (cl @p_95) :rule resolution :premises (t50 t51 t56) :args (@p_93 @p_8 @p_94 @p_8))
(step t58 (cl @p_88 @p_86) :rule or :premises (t57))
(anchor :step t59 :args ((:= (?v2 S_v_tm_s_v_tm_prod$) @p_27)))
(step t59.t0 (cl (= (or (not (member$a ?v2 @p_14)) (member$a ?v2 @p_15)) @p_88)) :rule refl)
(step t59 (cl @p_99) :rule sko_forall)
(step t60 (cl @p_101 @p_103 @p_13) :rule or_pos)
(step t61 (cl @p_104) :rule forall_inst :args (@p_14 @p_15))
(anchor :step t62 :args ((?v0 S_v_tm_s_v_tm_prod_set$) (?v1 S_v_tm_s_v_tm_prod_set$) (:= (?v0 S_v_tm_s_v_tm_prod_set$) ?v0) (:= (?v1 S_v_tm_s_v_tm_prod_set$) ?v1)))
(step t62.t0 (cl (= (forall ((?v2 S_v_tm_s_v_tm_prod$)) (=> (member$a ?v2 ?v0) (member$a ?v2 ?v1))) (forall ((?v2 S_v_tm_s_v_tm_prod$)) (or (not (member$a ?v2 ?v0)) (member$a ?v2 ?v1))))) :rule hole)
(step t62.t1 (cl (= (less_eq$a ?v0 ?v1) (less_eq$a ?v0 ?v1))) :rule refl)
(step t62.t2 (cl (= (=> (forall ((?v2 S_v_tm_s_v_tm_prod$)) (=> (member$a ?v2 ?v0) (member$a ?v2 ?v1))) (less_eq$a ?v0 ?v1)) (=> (forall ((?v2 S_v_tm_s_v_tm_prod$)) (or (not (member$a ?v2 ?v0)) (member$a ?v2 ?v1))) (less_eq$a ?v0 ?v1)))) :rule cong :premises (t62.t0 t62.t1))
(step t62 (cl @p_107) :rule bind)
(step t63 (cl @p_109) :rule hole)
(step t64 (cl @p_110) :rule refl)
(step t65 (cl @p_112) :rule refl)
(anchor :step t66 :args ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$) (:= (?v3 S_v_tm$) ?v3) (:= (?v4 S_v_tm$) ?v4) (:= (?v5 S_v_tm$) ?v5)))
(step t66.t0 (cl (= (= @p_27 (pair$ ?v3 ?v5)) (= (pair$ ?v3 ?v5) @p_27))) :rule hole)
(step t66.t1 (cl (= (not (= @p_27 (pair$ ?v3 ?v5))) (not (= (pair$ ?v3 ?v5) @p_27)))) :rule cong :premises (t66.t0))
(step t66.t2 (cl (= (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v3 ?v4) @p_15)))) :rule refl)
(step t66.t3 (cl (= (not (member$a (pair$ ?v4 ?v5) @p_17)) (not (member$a (pair$ ?v4 ?v5) @p_17)))) :rule refl)
(step t66.t4 (cl (= (or (not (= @p_27 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17))) (or (not (= (pair$ ?v3 ?v5) @p_27)) (not (member$a (pair$ ?v3 ?v4) @p_15)) (not (member$a (pair$ ?v4 ?v5) @p_17))))) :rule cong :premises (t66.t1 t66.t2 t66.t3))
(step t66 (cl @p_113) :rule bind)
(step t67 (cl @p_115) :rule cong :premises (t66))
(step t68 (cl @p_117) :rule cong :premises (t65 t67))
(step t69 (cl @p_119 @p_123 @p_122) :rule equiv_pos2)
(step t70 (cl @p_120) :rule cong :premises (t64 t68))
(step t71 (cl @p_121 @p_111) :rule implies_neg1)
(anchor :step t72)
(assume t72.a0 @p_111)
(step t72.t0 (cl @p_124) :rule forall_inst :args (@p_27 @p_15 @p_17))
(anchor :step t72.t1 :args ((?v0 S_v_tm_s_v_tm_prod$) (?v1 S_v_tm_s_v_tm_prod_set$) (?v2 S_v_tm_s_v_tm_prod_set$) (:= (?v0 S_v_tm_s_v_tm_prod$) ?v0) (:= (?v1 S_v_tm_s_v_tm_prod_set$) ?v1) (:= (?v2 S_v_tm_s_v_tm_prod_set$) ?v2)))
(step t72.t1.t0 (cl (= (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) @p_8) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)))))) :rule hole)
(step t72.t1.t1 (cl (= (member$a ?v0 (relcomp$ ?v1 ?v2)) (member$a ?v0 (relcomp$ ?v1 ?v2)))) :rule refl)
(anchor :step t72.t1.t2 :args ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$) (:= (?v3 S_v_tm$) ?v3) (:= (?v4 S_v_tm$) ?v4) (:= (?v5 S_v_tm$) ?v5)))
(step t72.t1.t2.t0 (cl (= (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8) (not (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule hole)
(step t72.t1.t2.t1 (cl (= (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) :rule hole)
(step t72.t1.t2.t2 (cl (= (not (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))))) :rule cong :premises (t72.t1.t2.t1))
(step t72.t1.t2.t3 (cl (= (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))))) :rule trans :premises (t72.t1.t2.t0 t72.t1.t2.t2))
(step t72.t1.t2 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule bind)
(step t72.t1.t3 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule hole)
(step t72.t1.t4 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule trans :premises (t72.t1.t2 t72.t1.t3))
(step t72.t1.t5 (cl (= (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2))))))) :rule cong :premises (t72.t1.t1 t72.t1.t4))
(step t72.t1.t6 (cl (= (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)))) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))))) :rule cong :premises (t72.t1.t5))
(step t72.t1.t7 (cl (= (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) @p_8) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))))) :rule trans :premises (t72.t1.t0 t72.t1.t6))
(step t72.t1 (cl @p_126) :rule bind)
(step t72.t2 (cl @p_128) :rule hole)
(step t72.t3 (cl @p_125 @p_118) :rule or :premises (t72.t0))
(step t72.t4 (cl @p_129 @p_131 @p_111) :rule equiv_pos2)
(step t72.t5 (cl @p_130) :rule trans :premises (t72.t1 t72.t2))
(step t72.t6 (cl @p_111) :rule resolution :premises (t72.t4 t72.t5 a6) :args (@p_130 @p_8 @p_7 @p_8))
(step t72.t7 (cl @p_118) :rule resolution :premises (t72.t3 t72.t6) :args (@p_111 @p_8))
(step t72 (cl @p_125 @p_118) :rule subproof :discharge (t72.a0))
(step t73 (cl @p_121 @p_118) :rule resolution :premises (t71 t72) :args (@p_111 @p_41))
(step t74 (cl @p_121 @p_132) :rule implies_neg2)
(step t75 (cl @p_121 @p_121) :rule resolution :premises (t73 t74) :args (@p_118 @p_41))
(step t76 (cl @p_121) :rule contraction :premises (t75))
(step t77 (cl @p_122) :rule resolution :premises (t69 t70 t76) :args (@p_120 @p_8 @p_121 @p_8))
(anchor :step t78 :args ((?v0 S_v_tm_s_v_tm_prod$) (?v1 S_v_tm_s_v_tm_prod_set$) (?v2 S_v_tm_s_v_tm_prod_set$) (:= (?v0 S_v_tm_s_v_tm_prod$) ?v0) (:= (?v1 S_v_tm_s_v_tm_prod_set$) ?v1) (:= (?v2 S_v_tm_s_v_tm_prod_set$) ?v2)))
(step t78.t0 (cl (= (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) @p_8) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)))))) :rule hole)
(step t78.t1 (cl (= (member$a ?v0 (relcomp$ ?v1 ?v2)) (member$a ?v0 (relcomp$ ?v1 ?v2)))) :rule refl)
(anchor :step t78.t2 :args ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$) (:= (?v3 S_v_tm$) ?v3) (:= (?v4 S_v_tm$) ?v4) (:= (?v5 S_v_tm$) ?v5)))
(step t78.t2.t0 (cl (= (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8) (not (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule hole)
(step t78.t2.t1 (cl (= (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) :rule hole)
(step t78.t2.t2 (cl (= (not (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))))) :rule cong :premises (t78.t2.t1))
(step t78.t2.t3 (cl (= (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))))) :rule trans :premises (t78.t2.t0 t78.t2.t2))
(step t78.t2 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule bind)
(step t78.t3 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (not (and (= ?v0 (pair$ ?v3 ?v5)) (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2)))) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule hole)
(step t78.t4 (cl (= (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))) :rule trans :premises (t78.t2 t78.t3))
(step t78.t5 (cl (= (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2))))))) :rule cong :premises (t78.t1 t78.t4))
(step t78.t6 (cl (= (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8)))) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))))) :rule cong :premises (t78.t5))
(step t78.t7 (cl (= (=> (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (=> (and (= ?v0 (pair$ ?v3 ?v5)) (and (member$a (pair$ ?v3 ?v4) ?v1) (member$a (pair$ ?v4 ?v5) ?v2))) @p_8))) @p_8) (not (and (member$a ?v0 (relcomp$ ?v1 ?v2)) (forall ((?v3 S_v_tm$) (?v4 S_v_tm$) (?v5 S_v_tm$)) (or (not (= ?v0 (pair$ ?v3 ?v5))) (not (member$a (pair$ ?v3 ?v4) ?v1)) (not (member$a (pair$ ?v4 ?v5) ?v2)))))))) :rule trans :premises (t78.t0 t78.t6))
(step t78 (cl @p_126) :rule bind)
(step t79 (cl @p_128) :rule hole)
(step t80 (cl @p_64 @p_62) :rule reordering :premises (t20))
(step t81 (cl @p_75 @p_72) :rule or :premises (t46))
(step t82 (cl @p_85 @p_74 @p_83) :rule reordering :premises (t47))
(step t83 (cl @p_86 @p_88) :rule reordering :premises (t58))
(step t84 (cl @p_134 @p_136 @p_97) :rule equiv_pos2)
(step t85 (cl @p_135) :rule cong :premises (t59))
(step t86 (cl @p_13 @p_103 @p_101) :rule reordering :premises (t60))
(step t87 (cl @p_105 @p_102) :rule or :premises (t61))
(step t88 (cl @p_137 @p_139 @p_106) :rule equiv_pos2)
(step t89 (cl @p_138) :rule trans :premises (t62 t63))
(step t90 (cl @p_106) :rule resolution :premises (t88 t89 a9) :args (@p_138 @p_8 @p_11 @p_8))
(step t91 (cl @p_102) :rule resolution :premises (t87 t90) :args (@p_106 @p_8))
(step t92 (cl @p_103) :rule resolution :premises (t86 a10 t91) :args (@p_13 @p_41 @p_102 @p_8))
(step t93 (cl @p_97) :rule resolution :premises (t84 t85 t92) :args (@p_135 @p_8 @p_103 @p_8))
(step t94 (cl @p_86) :rule resolution :premises (t83 t93) :args (@p_88 @p_41))
(step t95 (cl @p_125 @p_84) :rule implies :premises (t77))
(step t96 (cl @p_129 @p_131 @p_111) :rule equiv_pos2)
(step t97 (cl @p_130) :rule trans :premises (t78 t79))
(step t98 (cl @p_111) :rule resolution :premises (t96 t97 a6) :args (@p_130 @p_8 @p_7 @p_8))
(step t99 (cl @p_84) :rule resolution :premises (t95 t98) :args (@p_111 @p_8))
(step t100 (cl @p_74) :rule resolution :premises (t82 t94 t99) :args (@p_86 @p_8 @p_84 @p_8))
(step t101 (cl @p_72) :rule resolution :premises (t81 t100) :args (@p_75 @p_41))
(step t102 (cl @p_64) :rule resolution :premises (t80 t101) :args (@p_62 @p_41))
(step t103 (cl (= (! (= (! (fun_app$g @p_16 @p_27) :named @p_144) @p_8) :named @p_143) (! (not @p_144) :named @p_145))) :rule equiv_simplify)
(step t104 (cl (! (not (! (= @p_89 @p_144) :named @p_147)) :named @p_146) @p_89 @p_145) :rule equiv_pos1)
(step t105 (cl (or (! (not @p_6) :named @p_148) @p_147)) :rule forall_inst :args (@p_27 @p_16))
(step t106 (cl (= @p_16 @p_16)) :rule refl)
(step t107 (cl @p_76) :rule symm :premises (t102))
(step t108 (cl (= @p_57 @p_144)) :rule cong :premises (t106 t107))
(step t109 (cl @p_143 (! (not @p_145) :named @p_149)) :rule equiv2 :premises (t103))
(step t110 (cl (not @p_149) @p_144) :rule not_not)
(step t111 (cl @p_143 @p_144) :rule resolution :premises (t109 t110) :args (@p_149 @p_41))
(step t112 (cl @p_89 @p_145 @p_146) :rule reordering :premises (t104))
(step t113 (cl @p_88 (! (not @p_89) :named @p_150)) :rule or_neg :args (1))
(step t114 (cl @p_150) :rule resolution :premises (t113 t93) :args (@p_88 @p_41))
(step t115 (cl @p_148 @p_147) :rule or :premises (t105))
(step t116 (cl @p_147) :rule resolution :premises (t115 a5) :args (@p_6 @p_8))
(step t117 (cl @p_145) :rule resolution :premises (t112 t114 t116) :args (@p_89 @p_41 @p_147 @p_8))
(step t118 (cl @p_143) :rule resolution :premises (t111 t117) :args (@p_144 @p_41))
(step t119 (cl (or (! (not @p_9) :named @p_151) @p_56)) :rule forall_inst :args (uu$ @p_26 @p_25))
(step t120 (cl (or (! (not (! (forall ((?v0 S_v_tm$) (?v1 S_v_tm$)) (= (fun_app$ (fun_app$a uu$ ?v0) ?v1) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0))))) :named @p_153)) :named @p_152) @p_44)) :rule forall_inst :args (@p_26 @p_25))
(anchor :step t121 :args ((?v0 S_v_tm$) (?v1 S_v_tm$) (:= (?v0 S_v_tm$) ?v0) (:= (?v1 S_v_tm$) ?v1)))
(step t121.t0 (cl (= (fun_app$ (fun_app$a uu$ ?v0) ?v1) (fun_app$ (fun_app$a uu$ ?v0) ?v1))) :rule refl)
(step t121.t1 (cl (= (and (= (vars$ ?v1) bot$) (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))))))) (and (= (vars$ ?v1) bot$) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))))))) :rule hole)
(step t121.t2 (cl (= (= (vars$ ?v1) bot$) (= bot$ (vars$ ?v1)))) :rule hole)
(step t121.t3 (cl (= (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)))) :rule refl)
(anchor :step t121.t4 :args ((?v2 S_v_tm$) (?v3 S_v_tm$) (:= (?v2 S_v_tm$) ?v2) (:= (?v3 S_v_tm$) ?v3)))
(step t121.t4.t0 (cl (= (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))) :rule hole)
(step t121.t4 (cl (= (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))))) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))))) :rule bind)
(step t121.t5 (cl (= (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))) (not (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (not (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))))))) :rule hole)
(step t121.t6 (cl (= (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (not (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))) (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (or (not (= ?v1 ?v2)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))))) :rule hole)
(step t121.t7 (cl (= (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (or (not (= ?v1 ?v2)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))) (forall ((?v3 S_v_tm$)) (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3))))))) :rule hole)
(anchor :step t121.t8 :args ((?v3 S_v_tm$) (:= (?v3 S_v_tm$) ?v3)))
(step t121.t8.t0 (cl (= (= ?v1 ?v1) @p_41)) :rule hole)
(step t121.t8.t1 (cl (= (not (= ?v1 ?v1)) (! (not @p_41) :named @p_154))) :rule cong :premises (t121.t8.t0))
(step t121.t8.t2 (cl (! (= @p_154 @p_8) :named @p_155)) :rule hole)
(step t121.t8.t3 (cl (= (not (= ?v1 ?v1)) @p_8)) :rule trans :premises (t121.t8.t1 t121.t8.t2))
(step t121.t8.t4 (cl (= (not (= ?v0 ?v3)) (not (= ?v0 ?v3)))) :rule refl)
(step t121.t8.t5 (cl (= (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))))) :rule refl)
(step t121.t8.t6 (cl (= (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3))))) :rule refl)
(step t121.t8.t7 (cl (= (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3)))) (or @p_8 (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3)))))) :rule cong :premises (t121.t8.t3 t121.t8.t4 t121.t8.t5 t121.t8.t6))
(step t121.t8.t8 (cl (= (or @p_8 (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3)))) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3)))))) :rule hole)
(step t121.t8.t9 (cl (= (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3)))) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3)))))) :rule trans :premises (t121.t8.t7 t121.t8.t8))
(step t121.t8 (cl (= (forall ((?v3 S_v_tm$)) (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3))))) (forall ((?v3 S_v_tm$)) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3))))))) :rule bind)
(step t121.t9 (cl (= (forall ((?v3 S_v_tm$)) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3))))) (or (not (= ?v0 ?v0)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))) :rule hole)
(step t121.t10 (cl (= (= ?v0 ?v0) @p_41)) :rule hole)
(step t121.t11 (cl (= (not (= ?v0 ?v0)) @p_154)) :rule cong :premises (t121.t10))
(step t121.t12 (cl @p_155) :rule hole)
(step t121.t13 (cl (= (not (= ?v0 ?v0)) @p_8)) :rule trans :premises (t121.t11 t121.t12))
(step t121.t14 (cl (= (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))))) :rule refl)
(step t121.t15 (cl (= (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0))))) :rule refl)
(step t121.t16 (cl (= (or (not (= ?v0 ?v0)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))) (or @p_8 (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))) :rule cong :premises (t121.t13 t121.t14 t121.t15))
(step t121.t17 (cl (= (or @p_8 (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))) :rule hole)
(step t121.t18 (cl (= (or (not (= ?v0 ?v0)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))) :rule trans :premises (t121.t16 t121.t17))
(step t121.t19 (cl (= (forall ((?v3 S_v_tm$)) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))) :rule trans :premises (t121.t9 t121.t18))
(step t121.t20 (cl (= (forall ((?v3 S_v_tm$)) (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v3))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))) :rule trans :premises (t121.t8 t121.t19))
(step t121.t21 (cl (= (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (or (not (= ?v1 ?v2)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))) :rule trans :premises (t121.t7 t121.t20))
(step t121.t22 (cl (= (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (not (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))) :rule trans :premises (t121.t6 t121.t21))
(step t121.t23 (cl (= (not (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (not (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))))) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0))))))) :rule cong :premises (t121.t22))
(step t121.t24 (cl (= (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0))))))) :rule trans :premises (t121.t5 t121.t23))
(step t121.t25 (cl (= (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))))) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0))))))) :rule trans :premises (t121.t4 t121.t24))
(step t121.t26 (cl (= (and (= (vars$ ?v1) bot$) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))))) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))))) :rule cong :premises (t121.t2 t121.t3 t121.t25))
(step t121.t27 (cl (= (and (= (vars$ ?v1) bot$) (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3)))))))) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))))) :rule trans :premises (t121.t1 t121.t26))
(step t121.t28 (cl (= (= (fun_app$ (fun_app$a uu$ ?v0) ?v1) (and (= (vars$ ?v1) bot$) (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v2) (head$ ?v3))))))))) (= (fun_app$ (fun_app$a uu$ ?v0) ?v1) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0))))))))) :rule cong :premises (t121.t0 t121.t27))
(step t121 (cl (= @p_0 (! (forall ((?v0 S_v_tm$) (?v1 S_v_tm$)) (= (fun_app$ (fun_app$a uu$ ?v0) ?v1) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (gt_hd$ gt_sym$ ground_heads_var$ (head$ ?v1) (head$ ?v0)))))))) :named @p_156))) :rule bind)
(step t122 (cl (= @p_156 @p_153)) :rule hole)
(step t123 (cl (! (not (! (and @p_48 @p_23 (! (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ @p_24 @p_32) :named @p_159) (! (= @p_52 (! (head$ @p_31) :named @p_161)) :named @p_160) (! (forall ((?v4 S$)) (or (not (member$ ?v4 (! (ground_heads$ ground_heads_var$ @p_52) :named @p_163))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (! (args$ @p_25) :named @p_164)) (! (args$ @p_31) :named @p_165)))) :named @p_162)) :named @p_158)) :named @p_157) @p_48) :rule and_pos :args (0))
(step t124 (cl (! (not (! (= (! (fun_app$ (! (fun_app$a uua$ @p_31) :named @p_169) @p_25) :named @p_168) @p_158) :named @p_167)) :named @p_166) (! (not @p_168) :named @p_170) @p_158) :rule equiv_pos2)
(step t125 (cl (! (not (! (= (! (fun_app$g @p_18 @p_67) :named @p_173) @p_168) :named @p_172)) :named @p_171) (! (not @p_173) :named @p_174) @p_168) :rule equiv_pos2)
(step t126 (cl (! (not (! (= @p_66 @p_173) :named @p_176)) :named @p_175) @p_65 @p_173) :rule equiv_pos2)
(step t127 (cl (= (! (not @p_65) :named @p_177) @p_66)) :rule hole)
(step t128 (cl (not (! (= (! (or @p_62 @p_177) :named @p_179) (! (or @p_62 @p_66) :named @p_180)) :named @p_178)) (not @p_179) @p_180) :rule equiv_pos2)
(step t129 (cl @p_178) :rule cong :premises (t10 t127))
(step t130 (cl @p_62 @p_177) :rule or_neg :args (2))
(step t131 (cl @p_179 @p_72) :rule or_neg :args (0))
(step t132 (cl @p_179 (not @p_177)) :rule or_neg :args (1))
(step t133 (cl @p_179 @p_179) :rule resolution :premises (t130 t131 t132) :args (@p_62 @p_41 @p_177 @p_41))
(step t134 (cl @p_179) :rule contraction :premises (t133))
(step t135 (cl @p_180) :rule resolution :premises (t128 t129 t134) :args (@p_178 @p_8 @p_179 @p_8))
(step t136 (cl @p_62 @p_66) :rule or :premises (t135))
(step t137 (cl (or @p_148 @p_176)) :rule forall_inst :args (@p_67 @p_18))
(step t138 (cl (or @p_151 @p_172)) :rule forall_inst :args (uua$ @p_31 @p_25))
(step t139 (cl (or (! (not (! (forall ((?v0 S_v_tm$) (?v1 S_v_tm$)) (= (fun_app$ (fun_app$a uua$ ?v0) ?v1) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (= (head$ ?v1) (head$ ?v0)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0))))))) :named @p_182)) :named @p_181) @p_167)) :rule forall_inst :args (@p_31 @p_25))
(anchor :step t140 :args ((?v0 S_v_tm$) (?v1 S_v_tm$) (:= (?v0 S_v_tm$) ?v0) (:= (?v1 S_v_tm$) ?v1)))
(step t140.t0 (cl (= (fun_app$ (fun_app$a uua$ ?v0) ?v1) (fun_app$ (fun_app$a uua$ ?v0) ?v1))) :rule refl)
(step t140.t1 (cl (= (and (= (vars$ ?v1) bot$) (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))))))) (and (= (vars$ ?v1) bot$) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))))))))) :rule hole)
(step t140.t2 (cl (= (= (vars$ ?v1) bot$) (= bot$ (vars$ ?v1)))) :rule hole)
(step t140.t3 (cl (= (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)))) :rule refl)
(anchor :step t140.t4 :args ((?v2 S_v_tm$) (?v3 S_v_tm$) (:= (?v2 S_v_tm$) ?v2) (:= (?v3 S_v_tm$) ?v3)))
(step t140.t4.t0 (cl (= (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))))) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))) :rule hole)
(step t140.t4.t1 (cl (= (= ?v1 ?v2) (= ?v1 ?v2))) :rule refl)
(step t140.t4.t2 (cl (= (= ?v0 ?v3) (= ?v0 ?v3))) :rule refl)
(step t140.t4.t3 (cl (= (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)))) :rule refl)
(step t140.t4.t4 (cl (= (= (head$ ?v2) (head$ ?v3)) (= (head$ ?v2) (head$ ?v3)))) :rule refl)
(step t140.t4.t5 (cl (= (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))) :rule hole)
(step t140.t4.t6 (cl (= (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))) :rule cong :premises (t140.t4.t1 t140.t4.t2 t140.t4.t3 t140.t4.t4 t140.t4.t5))
(step t140.t4.t7 (cl (= (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))))) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))) :rule trans :premises (t140.t4.t0 t140.t4.t6))
(step t140.t4 (cl (= (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))))) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))))) :rule bind)
(step t140.t5 (cl (= (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))) (not (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (not (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))))))) :rule hole)
(step t140.t6 (cl (= (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (not (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))) (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (or (not (= ?v1 ?v2)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v2) (head$ ?v3))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))))) :rule hole)
(step t140.t7 (cl (= (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (or (not (= ?v1 ?v2)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v2) (head$ ?v3))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))) (forall ((?v3 S_v_tm$)) (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v1) (head$ ?v3))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3))))))))) :rule hole)
(anchor :step t140.t8 :args ((?v3 S_v_tm$) (:= (?v3 S_v_tm$) ?v3)))
(step t140.t8.t0 (cl (= (= ?v1 ?v1) @p_41)) :rule hole)
(step t140.t8.t1 (cl (= (not (= ?v1 ?v1)) @p_154)) :rule cong :premises (t140.t8.t0))
(step t140.t8.t2 (cl @p_155) :rule hole)
(step t140.t8.t3 (cl (= (not (= ?v1 ?v1)) @p_8)) :rule trans :premises (t140.t8.t1 t140.t8.t2))
(step t140.t8.t4 (cl (= (not (= ?v0 ?v3)) (not (= ?v0 ?v3)))) :rule refl)
(step t140.t8.t5 (cl (= (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))))) :rule refl)
(step t140.t8.t6 (cl (= (= (head$ ?v1) (head$ ?v3)) (= (head$ ?v3) (head$ ?v1)))) :rule hole)
(step t140.t8.t7 (cl (= (not (= (head$ ?v1) (head$ ?v3))) (not (= (head$ ?v3) (head$ ?v1))))) :rule cong :premises (t140.t8.t6))
(step t140.t8.t8 (cl (= (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3))))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3))))))) :rule refl)
(step t140.t8.t9 (cl (= (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v1) (head$ ?v3))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3)))))) (or @p_8 (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v3) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3)))))))) :rule cong :premises (t140.t8.t3 t140.t8.t4 t140.t8.t5 t140.t8.t7 t140.t8.t8))
(step t140.t8.t10 (cl (= (or @p_8 (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v3) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3)))))) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v3) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3)))))))) :rule hole)
(step t140.t8.t11 (cl (= (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v1) (head$ ?v3))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3)))))) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v3) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3)))))))) :rule trans :premises (t140.t8.t9 t140.t8.t10))
(step t140.t8 (cl (= (forall ((?v3 S_v_tm$)) (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v1) (head$ ?v3))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3))))))) (forall ((?v3 S_v_tm$)) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v3) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3))))))))) :rule bind)
(step t140.t9 (cl (= (forall ((?v3 S_v_tm$)) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v3) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3))))))) (or (not (= ?v0 ?v0)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v0) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))) :rule hole)
(step t140.t10 (cl (= (= ?v0 ?v0) @p_41)) :rule hole)
(step t140.t11 (cl (= (not (= ?v0 ?v0)) @p_154)) :rule cong :premises (t140.t10))
(step t140.t12 (cl @p_155) :rule hole)
(step t140.t13 (cl (= (not (= ?v0 ?v0)) @p_8)) :rule trans :premises (t140.t11 t140.t12))
(step t140.t14 (cl (= (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))))) :rule refl)
(step t140.t15 (cl (= (= (head$ ?v0) (head$ ?v1)) (= (head$ ?v1) (head$ ?v0)))) :rule hole)
(step t140.t16 (cl (= (not (= (head$ ?v0) (head$ ?v1))) (not (= (head$ ?v1) (head$ ?v0))))) :rule cong :premises (t140.t15))
(step t140.t17 (cl (= (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0))))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0))))))) :rule refl)
(step t140.t18 (cl (= (or (not (= ?v0 ?v0)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v0) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))) (or @p_8 (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))) :rule cong :premises (t140.t13 t140.t14 t140.t16 t140.t17))
(step t140.t19 (cl (= (or @p_8 (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))) :rule hole)
(step t140.t20 (cl (= (or (not (= ?v0 ?v0)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v0) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))) :rule trans :premises (t140.t18 t140.t19))
(step t140.t21 (cl (= (forall ((?v3 S_v_tm$)) (or (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v3) (head$ ?v1))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3))))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))) :rule trans :premises (t140.t9 t140.t20))
(step t140.t22 (cl (= (forall ((?v3 S_v_tm$)) (or (not (= ?v1 ?v1)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v1) (head$ ?v3))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v3))))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))) :rule trans :premises (t140.t8 t140.t21))
(step t140.t23 (cl (= (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (or (not (= ?v1 ?v2)) (not (= ?v0 ?v3)) (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3))) (not (= (head$ ?v2) (head$ ?v3))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))) :rule trans :premises (t140.t7 t140.t22))
(step t140.t24 (cl (= (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (not (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))) (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))) :rule trans :premises (t140.t6 t140.t23))
(step t140.t25 (cl (= (not (forall ((?v2 S_v_tm$) (?v3 S_v_tm$)) (not (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))))) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0))))))))) :rule cong :premises (t140.t24))
(step t140.t26 (cl (= (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (= ?v0 ?v3) (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0))))))))) :rule trans :premises (t140.t5 t140.t25))
(step t140.t27 (cl (= (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))))) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0))))))))) :rule trans :premises (t140.t4 t140.t26))
(step t140.t28 (cl (= (and (= (vars$ ?v1) bot$) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))))))) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))))) :rule cong :premises (t140.t2 t140.t3 t140.t27))
(step t140.t29 (cl (= (and (= (vars$ ?v1) bot$) (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3))))))))))) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))))) :rule trans :premises (t140.t1 t140.t28))
(step t140.t30 (cl (= (= (fun_app$ (fun_app$a uua$ ?v0) ?v1) (and (= (vars$ ?v1) bot$) (and (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (exists ((?v2 S_v_tm$) (?v3 S_v_tm$)) (and (= ?v1 ?v2) (and (= ?v0 ?v3) (and (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v2) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v3)) (and (= (head$ ?v2) (head$ ?v3)) (forall ((?v4 S$)) (=> (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v2))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v2)) (args$ ?v3)))))))))))) (= (fun_app$ (fun_app$a uua$ ?v0) ?v1) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0))))))))))) :rule cong :premises (t140.t0 t140.t29))
(step t140 (cl (= @p_1 (! (forall ((?v0 S_v_tm$) (?v1 S_v_tm$)) (= (fun_app$ (fun_app$a uua$ ?v0) ?v1) (and (= bot$ (vars$ ?v1)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0)) (not (or (not (ge_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v1) (wt$ ground_heads_var$ delta$ arity_sym$ wt_sym$ coef_sym$ ?v0))) (not (= (head$ ?v1) (head$ ?v0))) (not (forall ((?v4 S$)) (or (not (member$ ?v4 (ground_heads$ ground_heads_var$ (head$ ?v1)))) (fun_app$b (fun_app$c (fun_app$d (fun_app$e extf$ ?v4) @p_2) (args$ ?v1)) (args$ ?v0)))))))))) :named @p_183))) :rule bind)
(step t141 (cl (= @p_183 @p_182)) :rule hole)
(step t142 (cl (! (not (! (or @p_22 (! (not (! (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ @p_32 @p_35) :named @p_187)) :named @p_186) @p_50) :named @p_185)) :named @p_184) @p_22 @p_186 @p_50) :rule or_pos)
(step t143 (cl (! (not (! (and (! (= bot$ (! (vars$ @p_31) :named @p_191)) :named @p_190) @p_187 @p_34 (! (gt_hd$ gt_sym$ ground_heads_var$ @p_161 @p_53) :named @p_192)) :named @p_189)) :named @p_188) @p_187) :rule and_pos :args (1))
(step t144 (cl (! (not (! (= (! (fun_app$ @p_46 @p_31) :named @p_195) @p_189) :named @p_194)) :named @p_193) (! (not @p_195) :named @p_196) @p_189) :rule equiv_pos2)
(step t145 (cl (! (not (! (= (! (fun_app$g @p_16 @p_30) :named @p_199) @p_195) :named @p_198)) :named @p_197) (! (not @p_199) :named @p_200) @p_195) :rule equiv_pos2)
(step t146 (cl (! (not (! (= @p_29 @p_199) :named @p_202)) :named @p_201) @p_28 @p_199) :rule equiv_pos2)
(step t147 (cl (= (! (not @p_28) :named @p_203) @p_29)) :rule hole)
(step t148 (cl (not (! (= (! (or @p_62 @p_203) :named @p_205) (! (or @p_62 @p_29) :named @p_206)) :named @p_204)) (not @p_205) @p_206) :rule equiv_pos2)
(step t149 (cl @p_204) :rule cong :premises (t10 t147))
(step t150 (cl @p_62 @p_203) :rule or_neg :args (1))
(step t151 (cl @p_205 @p_72) :rule or_neg :args (0))
(step t152 (cl @p_205 (not @p_203)) :rule or_neg :args (1))
(step t153 (cl @p_205 @p_205) :rule resolution :premises (t150 t151 t152) :args (@p_62 @p_41 @p_203 @p_41))
(step t154 (cl @p_205) :rule contraction :premises (t153))
(step t155 (cl @p_206) :rule resolution :premises (t148 t149 t154) :args (@p_204 @p_8 @p_205 @p_8))
(step t156 (cl @p_62 @p_29) :rule or :premises (t155))
(step t157 (cl (or @p_148 @p_202)) :rule forall_inst :args (@p_30 @p_16))
(step t158 (cl (or @p_151 @p_198)) :rule forall_inst :args (uu$ @p_26 @p_31))
(step t159 (cl (or @p_152 @p_194)) :rule forall_inst :args (@p_26 @p_31))
(step t160 (cl @p_157 @p_23) :rule and_pos :args (1))
(step t161 (cl (or (! (not (! (forall ((?v0 V_pvar_hmultiset_tpoly$) (?v1 V_pvar_hmultiset_tpoly$) (?v2 V_pvar_hmultiset_tpoly$)) (or (not (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v0 ?v1)) (not (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v1 ?v2)) (eq_tpoly$ ground_heads_var$ delta$ arity_sym$ wt_sym$ ?v0 ?v2))) :named @p_208)) :named @p_207) @p_185)) :rule forall_inst :args (@p_24 @p_32 @p_35))
(step t162 (cl (= (! (= @p_51 @p_41) :named @p_209) @p_51)) :rule equiv_simplify)
(step t163 (cl @p_157 @p_160) :rule and_pos :args (3))
(step t164 (cl @p_170 @p_158 @p_166) :rule reordering :premises (t124))
(step t165 (cl @p_174 @p_168 @p_171) :rule reordering :premises (t125))
(step t166 (cl @p_65 @p_173 @p_175) :rule reordering :premises (t126))
(step t167 (cl @p_66 @p_62) :rule reordering :premises (t136))
(step t168 (cl @p_66) :rule resolution :premises (t167 t101) :args (@p_62 @p_41))
(step t169 (cl @p_148 @p_176) :rule or :premises (t137))
(step t170 (cl @p_176) :rule resolution :premises (t169 a5) :args (@p_6 @p_8))
(step t171 (cl @p_173) :rule resolution :premises (t166 t168 t170) :args (@p_66 @p_8 @p_176 @p_8))
(step t172 (cl @p_151 @p_172) :rule or :premises (t138))
(step t173 (cl @p_172) :rule resolution :premises (t172 a7) :args (@p_9 @p_8))
(step t174 (cl @p_168) :rule resolution :premises (t165 t171 t173) :args (@p_173 @p_8 @p_172 @p_8))
(step t175 (cl @p_181 @p_167) :rule or :premises (t139))
(step t176 (cl (not (! (= @p_1 @p_182) :named @p_210)) (not @p_1) @p_182) :rule equiv_pos2)
(step t177 (cl @p_210) :rule trans :premises (t140 t141))
(step t178 (cl @p_182) :rule resolution :premises (t176 t177 a1) :args (@p_210 @p_8 @p_1 @p_8))
(step t179 (cl @p_167) :rule resolution :premises (t175 t178) :args (@p_182 @p_8))
(step t180 (cl @p_158) :rule resolution :premises (t164 t174 t179) :args (@p_168 @p_8 @p_167 @p_8))
(step t181 (cl @p_160 @p_157) :rule reordering :premises (t163))
(step t182 (cl @p_160) :rule resolution :premises (t181 t180) :args (@p_158 @p_8))
(step t183 (cl (= (! (= @p_192 @p_41) :named @p_211) @p_192)) :rule equiv_simplify)
(step t184 (cl @p_188 @p_192) :rule and_pos :args (3))
(step t185 (cl (= gt_sym$ gt_sym$)) :rule refl)
(step t186 (cl (= ground_heads_var$ ground_heads_var$)) :rule refl)
(step t187 (cl (= @p_53 @p_53)) :rule refl)
(step t188 (cl (= @p_51 @p_192)) :rule cong :premises (t185 t186 t182 t187))
(step t189 (cl (not (! (= @p_0 @p_153) :named @p_212)) (not @p_0) @p_153) :rule equiv_pos2)
(step t190 (cl @p_212) :rule trans :premises (t121 t122))
(step t191 (cl @p_153) :rule resolution :premises (t189 t190 a0) :args (@p_212 @p_8 @p_0 @p_8))
(step t192 (cl @p_196 @p_189 @p_193) :rule reordering :premises (t144))
(step t193 (cl @p_200 @p_195 @p_197) :rule reordering :premises (t145))
(step t194 (cl @p_28 @p_199 @p_201) :rule reordering :premises (t146))
(step t195 (cl @p_29 @p_62) :rule reordering :premises (t156))
(step t196 (cl @p_29) :rule resolution :premises (t195 t101) :args (@p_62 @p_41))
(step t197 (cl @p_148 @p_202) :rule or :premises (t157))
(step t198 (cl @p_202) :rule resolution :premises (t197 a5) :args (@p_6 @p_8))
(step t199 (cl @p_199) :rule resolution :premises (t194 t196 t198) :args (@p_29 @p_8 @p_202 @p_8))
(step t200 (cl @p_151 @p_198) :rule or :premises (t158))
(step t201 (cl @p_198) :rule resolution :premises (t200 a7) :args (@p_9 @p_8))
(step t202 (cl @p_195) :rule resolution :premises (t193 t199 t201) :args (@p_199 @p_8 @p_198 @p_8))
(step t203 (cl @p_152 @p_194) :rule or :premises (t159))
(step t204 (cl @p_194) :rule resolution :premises (t203 t191) :args (@p_153 @p_8))
(step t205 (cl @p_189) :rule resolution :premises (t192 t202 t204) :args (@p_195 @p_8 @p_194 @p_8))
(step t206 (cl @p_211 (not @p_192)) :rule equiv2 :premises (t183))
(step t207 (cl @p_192 @p_188) :rule reordering :premises (t184))
(step t208 (cl @p_192) :rule resolution :premises (t207 t205) :args (@p_189 @p_8))
(step t209 (cl @p_211) :rule resolution :premises (t206 t208) :args (@p_192 @p_8))
(step t210 (cl @p_188 @p_34) :rule and_pos :args (2))
(step t211 (cl @p_37 @p_21) :rule implies :premises (t5))
(step t212 (cl @p_36 @p_33 @p_22 @p_42) :rule reordering :premises (t6))
(step t213 (cl @p_47 (not @p_48) (not @p_50) (! (not @p_36) :named @p_213) (not @p_51)) :rule and_neg)
(step t214 (cl @p_45 @p_54 @p_43) :rule reordering :premises (t7))
(step t215 (cl @p_57 @p_59 @p_55) :rule reordering :premises (t8))
(step t216 (cl (not @p_60) @p_61) :rule equiv1 :premises (t9))
(step t217 (cl @p_60) :rule trans :premises (t108 t118))
(step t218 (cl @p_61) :rule resolution :premises (t216 t217) :args (@p_60 @p_8))
(step t219 (cl @p_151 @p_56) :rule or :premises (t119))
(step t220 (cl @p_59 @p_151) :rule resolution :premises (t215 t218 t219) :args (@p_57 @p_41 @p_56 @p_8))
(step t221 (cl @p_152 @p_44) :rule or :premises (t120))
(step t222 (cl @p_151 @p_54 @p_152) :rule resolution :premises (t214 t220 t221) :args (@p_45 @p_41 @p_44 @p_8))
(step t223 (cl @p_48 @p_157) :rule reordering :premises (t123))
(step t224 (cl @p_50 @p_186 @p_22 @p_184) :rule reordering :premises (t142))
(step t225 (cl @p_187 @p_188) :rule reordering :premises (t143))
(step t226 (cl @p_23 @p_157) :rule reordering :premises (t160))
(step t227 (cl @p_207 @p_185) :rule or :premises (t161))
(step t228 (cl (not (! (= @p_3 @p_208) :named @p_214)) (not @p_3) @p_208) :rule equiv_pos2)
(step t229 (cl @p_214) :rule hole)
(step t230 (cl @p_208) :rule resolution :premises (t228 t229 a2) :args (@p_214 @p_8 @p_3 @p_8))
(step t231 (cl @p_185) :rule resolution :premises (t227 t230) :args (@p_208 @p_8))
(step t232 (cl @p_50 @p_188 @p_22) :rule resolution :premises (t224 t225 t231) :args (@p_187 @p_8 @p_185 @p_8))
(step t233 (cl (not @p_209) @p_51) :rule equiv1 :premises (t162))
(step t234 (cl @p_209) :rule trans :premises (t188 t209))
(step t235 (cl @p_51) :rule resolution :premises (t233 t234) :args (@p_209 @p_8))
(step t236 (cl @p_22 @p_152 @p_157 @p_213 @p_151 @p_188) :rule resolution :premises (t213 t222 t223 t232 t235) :args (@p_47 @p_41 @p_48 @p_8 @p_50 @p_8 @p_51 @p_8))
(step t237 (cl @p_34 @p_188) :rule reordering :premises (t210))
(step t238 (cl @p_42 @p_188 @p_22 @p_152 @p_157 @p_151) :rule resolution :premises (t212 t236 t237) :args (@p_36 @p_41 @p_34 @p_8))
(step t239 (cl @p_38 @p_40 @p_20) :rule equiv_pos2)
(step t240 (cl @p_39) :rule hole)
(step t241 (cl @p_20) :rule resolution :premises (t239 t240 a3) :args (@p_39 @p_8 @p_4 @p_8))
(step t242 (cl @p_151 @p_188 @p_22 @p_152 @p_157) :rule resolution :premises (t211 t238 t241) :args (@p_21 @p_41 @p_20 @p_8))
(step t243 (cl) :rule resolution :premises (t242 t215 t205 t180 t191 t101 a7 a5) :args (@p_23 @p_8 @p_189 @p_8 @p_158 @p_8 @p_153 @p_8 @p_62 @p_41 @p_9 @p_8 @p_6 @p_8))
